```scala
import $file.hw5stdlib
import hw5stdlib._
```

Output:
```
Compiling hw5stdlib.sc
[32mimport [39m[36m$file.$        
[39m
[32mimport [39m[36mhw5stdlib._[39m
```

# Homework 5

In this assignment we will develop some helpful structures for building the next interpreter.


## Dictionaries

We will need to define the datatype $\texttt{Dictionary}$ before we write the first interpreter for lettuce. We'll begin by giving the definition below:

$$
\begin{aligned}
\textbf{Dictionary}\ A\ B\ ::&= \text{EmptyDict} \\
&\mid\ \text{Entry}\ A\ B\ (\textbf{Dictionary}\ A\ B)
\end{aligned}
$$

This will create dictionaries of key-value pairs which will be very useful for implementing contexts in the next homework assignment. The definition of this type in Scala is given below:

<!-- grade_id: cell-c777fe8e435f2d12, points: 0, grade: false, solution: true -->
```scala
sealed trait Dictionary[+A, +B]
case object EmptyDict extends Dictionary[Nothing, Nothing]
case class Entry[A, B](a1: A, b1: B, x: Dictionary[A, B]) extends Dictionary[A, B]
```

Output:
```
defined [32mtrait[39m [36mDictionary[39m
defined [32mobject[39m [36mEmptyDict[39m
defined [32mclass[39m [36mEntry[39m
```

For both of the following functions you will need to provide a parameter for an equality-checking function. This is so we can check if keys are equal to eachother

### Membership(In) (5pts)

Write a function called `isIn` which returns `True` if the Dictionary contains a given key and `False` if it does not.

<!-- grade_id: cell-23433950c52a6f5e, points: 0, grade: false, solution: true -->
```scala
def isIn[A, B](f: ((A, A)=> Bool), d: Dictionary[A, B], key: A): Bool = d match{
    case EmptyDict => False
    case Entry(x, y, dict) => or(f(x, key), isIn(f, dict, key))
}
```

Output:
```
defined [32mfunction[39m [36misIn[39m
```

<!-- grade_id: cell-4972a9f7ef5db900, points: 5, grade: true, solution: false -->
```scala
assert(isIn(nat_eq, Entry(five, True, EmptyDict), five) == True)
assert(isIn(nat_eq, Entry(eight, True, Entry(five, True, EmptyDict)), five) == True)
assert(isIn(nat_eq, Entry(five, True, EmptyDict), six) == False)
assert(isIn(nat_eq, EmptyDict, seven) == False)
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
```

### Lookup (5 pts)

Write a function that retrieves a value that corresponds to a key in a dictionary. If the key is not in the map then return Nothing.

$$
\texttt{lookup} : (a \rightarrow a \rightarrow \mathbb{B}) \rightarrow \texttt{Dict}\ a\ b \rightarrow a \rightarrow \texttt{Maybe}\ b
$$

<!-- grade_id: cell-2753a17919e5caa9, points: 0, grade: false, solution: true -->
```scala
def lookup[A, B](f: (A, A)=>Bool, dict: Dictionary[A, B], key: A): Maybe[B] = dict match
{
    case EmptyDict => None
    case Entry(x, y, d) => f(key, x) match{
        case True => Just(y)
        case False => lookup(f, d, key)
    }
    
}
```

Output:
```
defined [32mfunction[39m [36mlookup[39m
```

<!-- grade_id: cell-efcc8fd55430e3cc, points: 5, grade: true, solution: false -->
```scala
assert(lookup(nat_eq, Entry(five, "ham", EmptyDict), five) == Just("ham"))
assert(lookup(nat_eq, Entry(eight, seven, Entry(five, three, EmptyDict)), five) == Just(three))
assert(lookup(nat_eq, Entry(five, True, EmptyDict), six) == None)
assert(lookup(nat_eq, EmptyDict, seven) == None)
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
```

### Practice using these functions (5 pts)
#### Part A
Create the following dictionary as a scala value using this definition:
```
{
    five: True,
    six: False,
    seven: False
}
```
Place it in a variable named "`ans_3a`"

<!-- grade_id: cell-d081675d6d53a9cc, points: 0, grade: false, solution: true -->
```scala
val ans_3a:Dictionary[Nat, Bool] = Entry(five, True, Entry(six, False, Entry(seven, False, EmptyDict)))
```

Output:
```
[36mans_3a[39m: [32mDictionary[39m[[32mNat[39m, [32mBool[39m] = Entry(Succ(Succ(Succ(Succ(Succ(Zero))))),True,Entry(Succ(Succ(Succ(Succ(Succ(Succ(Zero)))))),False,Entry(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Zero))))))),False,EmptyDict)))
```

<!-- grade_id: cell-14d49877afa79bb3, points: 2, grade: true, solution: false -->
```scala
assert(lookup(nat_eq, ans_3a, five) == Just(True))
assert(lookup(nat_eq, ans_3a, six) == Just(False))
assert(lookup(nat_eq, ans_3a, seven) == Just(False))
passed(2)
```

Output:
```
*** Tests Passed (2 points) ***
```

#### Part B
For the following, use the super secret `secret_dict` variable defined below. Use the functions you just defined to complete the following:

1. Look up the value mapped by key `Positive(one)` and place it in variable `ans_3b1` 
2. Look up the value mapped by key `Positive(three)` and place it in variable `ans_3b2`
3. Check whether the key `Negative(three)` is in the map and place True or False in variable `ans_3b3`
4. Check whether the key `Negative(one)` is in the map and place True or False in variable `ans_3b4`


**Note**: For 1 and 2 remember to extract it from the maybe type, set the variable to the string `"does not exist"` if you get `None` instead of `Just(x)`.

```scala
val secret_dict = Entry(Positive(three), "abc",
                        Entry(Negative(one), "xyz",
                              Entry(Positive(one), "MIB", EmptyDict)))
```

Output:
```
[36msecret_dict[39m: [32mEntry[39m[[32mProduct[39m with [32mSerializable[39m with [32mInteger[39m, [32mString[39m] = Entry(Positive(Succ(Succ(Succ(Zero)))),abc,Entry(Negative(Succ(Zero)),xyz,Entry(Positive(Succ(Zero)),MIB,EmptyDict)))
```

<!-- grade_id: cell-41d83bd054b1372a, points: 0, grade: false, solution: true -->
```scala
val ans_3b1 = isIn(int_eq, secret_dict, Positive(one)) match{
    case False => "does not exist"
    case True => lookup(int_eq, secret_dict, Positive(one)) match{
        case Just(y) => y
    }
}

val ans_3b2 = isIn(int_eq, secret_dict, Positive(three)) match{
    case False => "does not exist"
    case True => lookup(int_eq, secret_dict, Positive(three)) match{
        case Just(y) => y
    }
}

val ans_3b3 = isIn(int_eq, secret_dict, Negative(three)) 

val ans_3b4 = isIn(int_eq, secret_dict, Negative(one)) 
```

Output:
```
[36mans_3b1[39m: [32mString[39m = [32m"MIB"[39m
[36mans_3b2[39m: [32mString[39m = [32m"abc"[39m
[36mans_3b3[39m: [32mBool[39m = False
[36mans_3b4[39m: [32mBool[39m = True
```

<!-- grade_id: cell-0ffa542d701f8f36, points: 3, grade: true, solution: false -->
```scala
// Hidden tests (3 pts)
```