```scala
import $file.hw7stdlib
import hw7stdlib._
```

Output:
```
Compiling /home/user/Downloads/Assignment7/hw7stdlib.sc[32mimport [39m[36m$file.$        
[39m
[32mimport [39m[36mhw7stdlib._[39m
```

# Homework 7

In this assignment we will implement functions and recursive let bindings to support recursive functions.

## Additional Syntax (10 pts)

Add the syntax as case classes necessary to implement functions. Below is the grammar to get you started along with the syntax we have defined so far:

$$\begin{array}{rcll}
\mathbf{Expr} & ::=\ & ... \\
 & | & FunDef(\textbf{String}, \textbf{Expr}) \\ 
 & | & FunCall(\textbf{Expr}, \textbf{Expr}) \\
 & | & LetRec(\textbf{String}, \textbf{String}, \textbf{Expr}, \textbf{Expr}) \\
\end{array}$$

**Note:** `Const` has been renamed to `Num`

<!-- grade_id: cell-6bb55f28265b3f32, points: 0, grade: false, solution: true -->
```scala
sealed trait Expr
case class Num(x : Integer) extends Expr
case class Bin(x : Bool) extends Expr
case class Ident(x : String) extends Expr
case class Plus(x : Expr, y : Expr) extends Expr
case class Minus(x : Expr, y : Expr) extends Expr
case class Mult(x : Expr, y : Expr) extends Expr
case class Pow(x : Expr, y : Expr) extends Expr
case class Neg(x : Expr) extends Expr
case class Eq(x : Expr, y : Expr) extends Expr
case class And(x : Expr, y : Expr) extends Expr
case class Or(x : Expr, y : Expr) extends Expr
case class IfThenElse(p : Expr, t : Expr, f : Expr) extends Expr
case class Let(id : String, x : Expr, y : Expr) extends Expr
case class FunDef(x: String, y: Expr) extends Expr
case class FunCall(x: Expr, y: Expr) extends Expr
case class LetRec(x: String, y: String, e1: Expr, e2: Expr) extends Expr
```

Output:
```
defined [32mtrait[39m [36mExpr[39m
defined [32mclass[39m [36mNum[39m
defined [32mclass[39m [36mBin[39m
defined [32mclass[39m [36mIdent[39m
defined [32mclass[39m [36mPlus[39m
defined [32mclass[39m [36mMinus[39m
defined [32mclass[39m [36mMult[39m
defined [32mclass[39m [36mPow[39m
defined [32mclass[39m [36mNeg[39m
defined [32mclass[39m [36mEq[39m
defined [32mclass[39m [36mAnd[39m
defined [32mclass[39m [36mOr[39m
defined [32mclass[39m [36mIfThenElse[39m
defined [32mclass[39m [36mLet[39m
defined [32mclass[39m [36mFunDef[39m
defined [32mclass[39m [36mFunCall[39m
defined [32mclass[39m [36mLetRec[39m
```

<!-- grade_id: cell-a865b9dd4be15e20, points: 4, grade: true, solution: false -->
```scala
val f = FunDef("x", Ident("x"))
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
[36mf[39m: [32mFunDef[39m = [33mFunDef[39m([32m"x"[39m, [33mIdent[39m([32m"x"[39m))
```

<!-- grade_id: cell-6d7a5b113b95a883, points: 3, grade: true, solution: false -->
```scala
val call = FunCall(f, Num(Positive(three)))
passed(3)
```

Output:
```
*** Tests Passed (3 points) ***
[36mcall[39m: [32mFunCall[39m = [33mFunCall[39m(
  [33mFunDef[39m([32m"x"[39m, [33mIdent[39m([32m"x"[39m)),
  [33mNum[39m([33mPositive[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero)))))
)
```

<!-- grade_id: cell-cde3f5ce0b49ec81, points: 3, grade: true, solution: false -->
```scala
val lrec = LetRec("f", "x", Plus(Ident("x"), Num(Positive(three))), call)
passed(3)
```

Output:
```
*** Tests Passed (3 points) ***
[36mlrec[39m: [32mLetRec[39m = [33mLetRec[39m(
  [32m"f"[39m,
  [32m"x"[39m,
  [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mNum[39m([33mPositive[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero)))))),
  [33mFunCall[39m([33mFunDef[39m([32m"x"[39m, [33mIdent[39m([32m"x"[39m)), [33mNum[39m([33mPositive[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero))))))
)
```

## Closures (5 pts)

Extend the defininition of $\mathbf{Value}$ to support closures:
$$\begin{array}{rcll}
\mathbf{Value} & ::=\ & ... \\
 & | & Closure(\textbf{String}, \textbf{Expr}, \textbf{Environment}) \\
\end{array}$$

### Note: 
This cell also contains the new definition for environments needed for recursion. It differs from using dictionaries in that:

1. it is non-polymorphic
2. The extend-rec constructor will be used to create closures for a let-rec binding. Where f is the key/name of the function and then x, e, and sigma represent the elements in the closure. We implement this functionality in the lookup function for environments

**When you implement the case for LetRec below you will want to use ExtendRec to extend the environment**

The following mathematical notation:

$$\sigma' = \sigma[x_\text{name} \mapsto \texttt{Closure}(x_\text{param}, e_\text{body}, \sigma')]$$

would be this in scala using the new definition:

```scala
val sigma_prime = ExtendRec(x_name, x_param, e_body, sigma)
```

<!-- grade_id: cell-232282dd59006346, points: 0, grade: false, solution: true -->
```scala
// Values
sealed trait Value
case class NumVal(x : Integer) extends Value
case class BinVal(x : Bool) extends Value
case object Error extends Value
case class Closure(x: String, y: Expr, z: Environment) extends Value

// New definition for environment for recursive funcs:
sealed trait Environment
case object EmptyEnv extends Environment 
case class Extend(k : String, v : Value, env : Environment) extends Environment
case class ExtendRec(f: String, x: String, e: Expr, sigma: Environment ) extends Environment
```

Output:
```
defined [32mtrait[39m [36mValue[39m
defined [32mclass[39m [36mNumVal[39m
defined [32mclass[39m [36mBinVal[39m
defined [32mobject[39m [36mError[39m
defined [32mclass[39m [36mClosure[39m
defined [32mtrait[39m [36mEnvironment[39m
defined [32mobject[39m [36mEmptyEnv[39m
defined [32mclass[39m [36mExtend[39m
defined [32mclass[39m [36mExtendRec[39m
```

<!-- grade_id: cell-5f8d075d8c1343ea, points: 5, grade: true, solution: false -->
```scala
val c = Closure("x", Ident("x"), EmptyEnv)
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
[36mc[39m: [32mClosure[39m = [33mClosure[39m([32m"x"[39m, [33mIdent[39m([32m"x"[39m), EmptyEnv)
```

### New lookup function
This function is used instead of the one for dictionaries when working with environments. You don't need to use it as the Ident case is filled out already.

```scala
def lookupEnv(sigma: Environment, x: String): Maybe[Value] = sigma match {
    case EmptyEnv => None
    case Extend(y, v, pi) => string_eq(y,x) match {
        case True => Just(v)
        case False => lookupEnv(pi, x)
    } 
    case ExtendRec(f, y, e, pi) => string_eq(x,f) match {
        case True => Just(Closure(y, e, sigma))
        case False => lookupEnv(pi, x)
    } 
}   
```

Output:
```
defined [32mfunction[39m [36mlookupEnv[39m
```

## Updating the interpreter (20 pts)

Add new cases to the interpreter to handle the new expressions we have added to the language. Once this is done you should be able to support both recursive and non-recursive functions. We have given you the skeleton to get started.

<!-- grade_id: cell-32d0d95205e2063f, points: 0, grade: false, solution: true -->
```scala
def eval(env : Environment, expr : Expr) : Value = expr match {
    case Num(n)    => NumVal(n)
    case Bin(p)    => BinVal(p)
    case Ident(id) => lookupEnv(env, id) match {
        case Just(v) => v
        case None    => Error
    }
    case Plus(e1, e2)  => eval_bin_arith(plus, env, e1, e2)
    case Minus(e1, e2) => eval_bin_arith(minus, env, e1, e2)
    case Mult(e1, e2)  => eval_bin_arith(mult, env, e1, e2)
    case Pow(e1, e2)   => (eval(env, e1), eval(env,e2)) match {
        case (NumVal(x), NumVal(Positive(n))) => NumVal(pow(x, n))
        case _ => Error
    }
    case Neg(e) => eval(env, e) match {
        case NumVal(x) => NumVal(negate(x))
        case BinVal(p) => BinVal(not(p))
        case Error     => Error
    }
    case Eq(e1, e2) => (eval(env, e1), eval(env, e2)) match {
        case (NumVal(x), NumVal(y)) => BinVal(int_eq(x,y))
        case (BinVal(p), BinVal(q)) => BinVal(bool_eq(p,q))
        case _                      => Error
    }
    case And(e1, e2) => eval_bin_bool(and, env, e1, e2)
    case Or(e1, e2)  => eval_bin_bool(or, env, e1, e2)
    case IfThenElse(p, e_t, e_f) => eval(env, p) match{
        case BinVal(True)  => eval(env, e_t)
        case BinVal(False) => eval(env, e_f)
        case _             => Error
    }
    case Let(id, df, body) => eval(env, df) match{
        case Error => Error
        case x     => {
            val new_env = Extend(id, x, env)
            eval(new_env , body)
        }
    }
    
    case FunDef(x, e_body) => Closure(x, e_body, env)
    
    case FunCall(e_f, e_arg) => eval(env, e_f) match
    {
        case Closure(x, ex, env1) => eval(env, e_arg) match
        {
            case Error => Error
            case x1 => {eval(Extend(x, x1, env1), ex)}
        }
        case _ => Error
    }
        
    
    case LetRec(x_name, x_param, e_body, e_in) => eval(ExtendRec(x_name, x_param, e_body, env), e_in) 
        
}

def eval_bin_arith( op : (Integer, Integer) => Integer
                  , env : Environment
                  , e1 : Expr
                  , e2 : Expr) : Value 
    = (eval(env, e1), eval(env, e2)) match{
        case (NumVal(x), NumVal(y)) => NumVal(op(x,y))
        case _ => Error
    }

def eval_bin_bool( op : (Bool, Bool) => Bool
                 , env : Environment
                 , e1 : Expr
                 , e2 : Expr) : Value 
    = (eval(env, e1), eval(env, e2)) match{
        case (BinVal(x), BinVal(y)) => BinVal(op(x,y))
        case _ => Error
    }
```

Output:
```
defined [32mfunction[39m [36meval[39m
defined [32mfunction[39m [36meval_bin_arith[39m
defined [32mfunction[39m [36meval_bin_bool[39m
```

<!-- grade_id: cell-c792093c303f5ff3, points: 5, grade: true, solution: false -->
```scala
//
// Tests Function definition
//

val x = Ident("x")
val f = Ident("f")

// function(x) x
val e_1 = FunDef("x", x)
assert(eval(EmptyEnv, e_1) == Closure("x", x, EmptyEnv))
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
[36mx[39m: [32mIdent[39m = [33mIdent[39m([32m"x"[39m)
[36mf[39m: [32mIdent[39m = [33mIdent[39m([32m"f"[39m)
[36me_1[39m: [32mFunDef[39m = [33mFunDef[39m([32m"x"[39m, [33mIdent[39m([32m"x"[39m))
```

<!-- grade_id: cell-9fe33b6093804635, points: 5, grade: true, solution: false -->
```scala
//
// Tests Function Calls
//

// (function(x) x)(2)
val e_2 = FunCall(e_1, Num(Positive(two)))
assert(eval(EmptyEnv, e_2) == NumVal(Positive(two)))
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
[36me_2[39m: [32mFunCall[39m = [33mFunCall[39m([33mFunDef[39m([32m"x"[39m, [33mIdent[39m([32m"x"[39m)), [33mNum[39m([33mPositive[39m([33mSucc[39m([33mSucc[39m(Zero)))))
```

<!-- grade_id: cell-4a098bd53eae5460, points: 5, grade: true, solution: false -->
```scala
//
// Tests Closed environments
//

// let y = 3 in
//   let f = function(x) y in
//     let y = 4 in
//       f(0)
val e_3 =
    Let("y", Num(Positive(three)),
        Let("f", FunDef("x", Ident("y")),
            Let("y", Num(Positive(four)),
                FunCall(f, Num(Positive(Zero)))
            )
        )
    )
assert(eval(EmptyEnv, e_3) == NumVal(Positive(three)))
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
[36me_3[39m: [32mLet[39m = [33mLet[39m(
  [32m"y"[39m,
  [33mNum[39m([33mPositive[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero))))),
  [33mLet[39m(
    [32m"f"[39m,
    [33mFunDef[39m([32m"x"[39m, [33mIdent[39m([32m"y"[39m)),
    [33mLet[39m(
      [32m"y"[39m,
      [33mNum[39m([33mPositive[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero)))))),
      [33mFunCall[39m([33mIdent[39m([32m"f"[39m), [33mNum[39m([33mPositive[39m(Zero)))
    )
  )
)
```

<!-- grade_id: cell-eba386ff1a243aa5, points: 5, grade: true, solution: false -->
```scala
//
// Tests Recursion
//

// if x == 0 then 1 else x * f(x - 1)
val factorial_body = IfThenElse(
    Eq(x, Num(Positive(Zero))),
    Num(Positive(one)),
    Mult(x, FunCall(f, Minus(x, Num(Positive(one)))))
)

// let rec f = function(x) if x == 0 then 1 else x * f(x - 1)
//     in f(3)
val e_1 = LetRec("f", "x", factorial_body, FunCall(f, Num(Positive(three))))

assert(eval(EmptyEnv, e_1) == NumVal(Positive(six)))
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
[36mfactorial_body[39m: [32mIfThenElse[39m = [33mIfThenElse[39m(
  [33mEq[39m([33mIdent[39m([32m"x"[39m), [33mNum[39m([33mPositive[39m(Zero))),
  [33mNum[39m([33mPositive[39m([33mSucc[39m(Zero))),
  [33mMult[39m(
    [33mIdent[39m([32m"x"[39m),
    [33mFunCall[39m([33mIdent[39m([32m"f"[39m), [33mMinus[39m([33mIdent[39m([32m"x"[39m), [33mNum[39m([33mPositive[39m([33mSucc[39m(Zero)))))
  )
)
[36me_1[39m: [32mLetRec[39m = [33mLetRec[39m(
  [32m"f"[39m,
  [32m"x"[39m,
  [33mIfThenElse[39m(
    [33mEq[39m([33mIdent[39m([32m"x"[39m), [33mNum[39m([33mPositive[39m(Zero))),
    [33mNum[39m([33mPositive[39m([33mSucc[39m(Zero))),
    [33mMult[39m(
      [33mIdent[39m([32m"x"[39m),
      [33mFunCall[39m([33mIdent[39m([32m"f"[39m), [33mMinus[39m([33mIdent[39m([32m"x"[39m), [33mNum[39m([33mPositive[39m([33mSucc[39m(Zero)))))
    )
  ),
  [33mFunCall[39m([33mIdent[39m(
```

```scala

```