```scala
import $file.hw3stdlib_new
import hw3stdlib_new._
```

Output:
```
[32mimport [39m[36m$file.$            
[39m
[32mimport [39m[36mhw3stdlib_new._[39m
```

# Homework 3

Due 9/22 at 11:59pm

## Submission Instructions
Upload only this .ipynb file to Canvas. Do not add anything to hw2stdlib since you can't submit it.

In this homework we will develop more higher order functions and learn how to use fold.

See [This link](https://www.notion.so/Guidelines-for-Programming-Homework-dbd25efa7bb24915ae6bcb06827fc5b6) for what is and isn't allowed in your code. 

## Problem 1 (5 Points)

Write the filter function. Remember that this should be polymorphic. So:

$$
filter : (a \rightarrow \mathbb{B}) \rightarrow List\ a \rightarrow List\ a
$$

Do not curry the parameters in the Scala version. If the predicate (The function we give it) is true for an element, then that element will be in the output list. If it's false leave it out.

<!-- grade_id: cell-9b2228c9dd99027a, points: 0, grade: false, solution: true -->
```scala
def filter[a](f: (a =>Bool), x: List[a]) : List[a] = x match {
    case Empty => Empty
    case Cons(z, zs) => f(z) match {
        case True => Cons(z, filter(f, zs))
        case False => filter(f, zs)
    }
}
```

Output:
```
defined [32mfunction[39m [36mfilter[39m
```

<!-- grade_id: cell-b64b76d103b58574, points: 4, grade: true, solution: false -->
```scala
assert(filter((n: Nat) => lte(n, two), Cons(three, Cons(two, Empty))) == Cons(two, Empty), 1)
assert(filter((n: Nat) => lte(n, two), Empty) == Empty, 2)
assert(filter((n: Nat) => lte(n, four), Cons(three, Cons(two, Empty))) == Cons(three, Cons(two, Empty)), 3)
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
```

<!-- grade_id: cell-0539a9bbd0c1b24e, points: 1, grade: true, solution: false -->
```scala
// HIDDEN TEST (1 pts)
```

### Problem 2 (5 Points)

Implement the same filter function using a `fold`. Name it `filterWithFold`. (Hint, take a look at the `append` and / or `reverse` functions in the standard library).

<!-- grade_id: cell-16f7380070a66d8e, points: 0, grade: false, solution: true -->
```scala
def filterWithFold[a](f: (a =>Bool), x: List[a]) : List[a] = 
    reverse(fold((y: a, b: List[a]) => f(y) match{
        case True => Cons(y, b)
        case False => b
    }, Empty, x))

```

Output:
```
defined [32mfunction[39m [36mfilterWithFold[39m
```

<!-- grade_id: cell-49c4b0530d9db3ca, points: 4, grade: true, solution: false -->
```scala
assert(filterWithFold((n: Nat) => lte(n, two), Cons(three, Cons(two, Empty))) == Cons(two, Empty), 1)
assert(filterWithFold((n: Nat) => lte(n, two), Empty) == Empty, 2)
assert(filterWithFold((n: Nat) => lte(n, four), Cons(three, Cons(two, Empty))) == Cons(three, Cons(two, Empty)), 3)
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
```

<!-- grade_id: cell-633a0587e8ffb49f, points: 1, grade: true, solution: false -->
```scala
// HIDDEN TEST (1 pts)
```

### Problem 3 (5 points)
Implement a function $$ifThenElse: \mathbb{B} \rightarrow a \rightarrow a \rightarrow a$$ which chooses either the first $A$ given if the bool is true or the second if it's false.

<!-- grade_id: cell-bcd9251068007fe3, points: 0, grade: false, solution: true -->
```scala
def ifThenElse[a](x: Bool)(a1: a)(a2: a) : a = x match {
    case True => a1
    case False => a2
}
```

Output:
```
defined [32mfunction[39m [36mifThenElse[39m
```

<!-- grade_id: cell-fc7f7449ac47fa14, points: 4, grade: true, solution: false -->
```scala
assert(ifThenElse(True)(one)(two) == one)
assert(ifThenElse(False)(one)(two) == two)
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
```

<!-- grade_id: cell-d4683268f4a95148, points: 1, grade: true, solution: false -->
```scala
// HIDDEN TEST (1 pts)
```

### Problem 4 (5 points)
Implement the `Maybe` type:
$$
\begin{align}
Maybe\ a\ :=\ \text{None}\ \mid\ \text{Just}\ a\
\end{align}
$$
Take a look at the definition of `List` in the stdlib as a starting point

<!-- grade_id: cell-9ee3e730ef837143, points: 0, grade: false, solution: true -->
```scala
sealed trait Maybe[+a]
case object None extends Maybe[Nothing]
case class Just[a](x: a) extends Maybe[a]
```

Output:
```
defined [32mtrait[39m [36mMaybe[39m
defined [32mobject[39m [36mNone[39m
defined [32mclass[39m [36mJust[39m
```

<!-- grade_id: cell-7c1e8b1d9b6f1079, points: 4, grade: true, solution: false -->
```scala
val mx: Maybe[Nat] = None
val my = Just(three)
(None: Maybe[Nat]) match {
    case None => two
    case Just(n) => three
}
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
[36mmx[39m: [32mMaybe[39m[[32mNat[39m] = None
[36mmy[39m: [32mJust[39m[[32mSucc[39m] = Just(Succ(Succ(Succ(Zero))))
[36mres29_2[39m: [32mSucc[39m = Succ(Succ(Zero))
```

<!-- grade_id: cell-8201540dbf3222f9, points: 1, grade: true, solution: false -->
```scala
// HIDDEN TEST (1 pts)
```

### Problem 5 (5 points)
Implement

$$
map: (a \rightarrow b) \rightarrow Maybe\ a \rightarrow Maybe\ b
$$

Don't curry the function in the Scala implementation. Similarly to lists, it should return `None` if given `None` and should return `Just(f(value))` if it contains a value.

<!-- grade_id: cell-2798ee2b3f2a50f6, points: 0, grade: false, solution: true -->
```scala
def map[a,b](f: (a => b), x: Maybe[a]) : Maybe[b] = x match {
    case None => None
    case Just(x) => Just(f(x))
}
```

Output:
```
defined [32mfunction[39m [36mmap[39m
```

<!-- grade_id: cell-722112b63bcac170, points: 4, grade: true, solution: false -->
```scala
assert(map(plus(_: Nat, four), None) == None)
assert(map(plus(_: Nat, four), Just(one)) == Just(five))
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
```

<!-- grade_id: cell-0ce38fa2255b8c9d, points: 1, grade: true, solution: false -->
```scala
// HIDDEN TEST (1 pts)
```